## Exercice 1
Ecrire une fonction **isPrime** qui teste si un nombre entier est premier

<!-- grade_id: cell-4e9946f987db9773, points: 0, grade: false, solution: true -->
```python
from math import sqrt
def isPrime(n):
    ### BEGIN SOLUTION
    if n==2 or n==3:
        return True
    if n%2 == 0 or n<=1:
        return False
    d=3
    while n%d!=0 and d <= sqrt(n):
        d+=2
    return n%d!=0
    ### END SOLUTION
```

Vérifications : La cellule ci-dessous doit s'exécuter sans erreur

<!-- grade_id: cell-a7acd19af3fad18b, points: 2, grade: true, solution: false -->
```python
assert isPrime(5) == True
assert isPrime(9) == False
### BEGIN HIDDEN TESTS
assert isPrime(0) == False
assert isPrime(1) == False
assert isPrime(2) == True
assert isPrime(3) == True
### END HIDDEN TESTS
```

## Exercice 2

En utilisant la fonction précédente, écrire une fonction **nbPrimesBetween**
- prenant en paramètres deux nombres $a$ et $b$
- renvoyant le nombre de nombres premiers entre $a$ et $b$

<!-- grade_id: cell-9571a0f46c204c37, points: 0, grade: false, solution: true -->
```python
def nbPrimesBetween(a,b):
    ### BEGIN SOLUTION
    return len([i for i in range(a, b+1) if isPrime(i)])
    ### END SOLUTION
```

Vérifications : La cellule ci-dessous doit s'exécuter sans erreur

<!-- grade_id: cell-3dfa0207efa01593, points: 0, grade: true, solution: false -->
```python
assert nbPrimesBetween(100,1000) == 143
### BEGIN HIDDEN TESTS
assert nbPrimesBetween(0,100) == 25
assert nbPrimesBetween(1000,100000) == 9424
### END HIDDEN TESTS
```